Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you :
- Fill in your name
- Fill any place that says `YOUR CODE HERE` (you will be graded based on this code, any extra cells you code will not be graded)
- Delete  `NotImplementedError()` exception once you have completed your implementation

```python
NAME = "NGUYEN THI THU TRANG"
```

---

# **PRIORITY QUEUE**

## **Problem 1: Min Binary Heap (4)**

Implement `MinHeap` class which follows the Binary Min Heap data structure. You can use `Item` class given below:

```python
class Item:
    "Item in Heap."

    def __init__(self, key, value):
        self.key = key
        self.val = value

    def __lt__(self, other):
        return self.key < other.key

    def __gt__(self, other):
        return self.key > other.key

    def __eq__(self, other):
        return self.key == other.key

    def __str__(self):
        return str((self.key, self.val))
```

Your task is complete the following methods in the class `MinHeap` below:



<!-- grade_id: cell-61f25814fd641a91, points: 0, grade: false, solution: true -->
```python
class MinHeap:
    def __init__(self):
        " Initialize an empty MinHeap."
        self._data = [] 

    def __len__(self):
        "Returns the number of items in the heap."
        return len(self._data)

    def is_empty(self):
        "Checks if the heap is empty. "
        return len(self._data) == 0

    def _parent(self, k: int):
        "Returns the index of the parent of the item at index k."
        return (k - 1) // 2

    def _left(self, k):
        "Returns the index of the left child of the item at index k."
        return k * 2 + 1

    def _right(self, k):
        "Returns the index of the right child of the item at index k."
        return k * 2 + 2

    def _has_left(self, k):
        "Checks if the item at index k has a left child."
        if self._data[self._left(k)] is None:
            return False
        return True

    def _has_right(self, k):
        "Checks if the item at index k has a right child."
        if self._data[self._right(k)] is None:
            return False
        return True

    def add(self, item: Item):
        """
        Adds a new item into the heap.
        Note: If the new item’s key has already existed in the heap, raise appropriate exception.
        """
        for existing_item in self._data:
            if existing_item.key == item.key:
                raise ValueError(f"Key {item.key} already exists in the heap.")

        self._data.append(item)

        index = len(self._data) - 1
        while index > 0:
            parent = self._parent(index)
            if self._data[index] < self._data[parent]:
                self._data[index], self._data[parent] = self._data[parent], self._data[index] 
                index = parent
            else:
                break

    def get_min(self):
        "Returns the item with minimum key in the heap."
        if self.is_empty():
            raise ValueError("Heap is empty.")
        return self._data[0]

    def remove_min(self):
        "Removes the item with minimum key in the heap."
        if self.is_empty():
            raise ValueError("Heap is empty.")

        self._data[0], self._data[-1] = self._data[-1], self._data[0]
        min_item = self._data.pop()

        index = 0
        while True:
            left = self._left(index)
            right = self._right(index)
            smallest = index

            if left < len(self._data) and self._data[left] < self._data[smallest]:
                smallest = left
            if right < len(self._data) and self._data[right] < self._data[smallest]:
                smallest = right
            
            if smallest == index:
                break

            self._data[index], self._data[smallest] = self._data[smallest], self._data[index]
            index = smallest
        
        return min_item

    def max_heap_sort(self):
        "Sorts the current min heap to a max heap."
        n = len(self._data)
        for index in range((n - 2) // 2, - 1, -1):
            i = index
            while True:
                left = self._left(i)
                right = self._right(i)
                largest = i

                if left < n and self._data[largest] < self._data[left]:
                    largest = left
                if right < n and self._data[largest] < self._data[right]:
                    largest = right
                
                if largest == i:
                    break

                self._data[i], self._data[largest] = self._data[largest], self._data[i]
                i = largest
                
    def __str__(self):
        "Returns a string representation of the heap."
        ans = []
        for item in self._data:
            ans.append(str(item))
        return "[" + ", ".join(ans) + "]"
    
    def __repr__(self):
        return self.__str__()

```

<!-- grade_id: cell-271a3446909a621c, points: 1, grade: true, solution: false -->
```python
# Initialize a MinHeap
heap = MinHeap()

# Test adding elements to the heap
for item in [(2, 'E'), (7, 'A'), (5, 'S'), (1, 'S'), (0, 'D'), (8, 'U'), (9, 'B'), (4, 'B'), (10, 'A')]:
    heap.add(Item(item[0], item[1]))
```

<!-- grade_id: cell-545b3dc78835408b, points: 1, grade: true, solution: false -->
```python
# Test heap property after adding elements
print(heap)
assert str(heap) == "[(0, 'D'), (1, 'S'), (5, 'S'), (4, 'B'), (2, 'E'), (8, 'U'), (9, 'B'), (7, 'A'), (10, 'A')]"
```

Output:
```
[(0, 'D'), (1, 'S'), (5, 'S'), (4, 'B'), (2, 'E'), (8, 'U'), (9, 'B'), (7, 'A'), (10, 'A')]
```

<!-- grade_id: cell-c4791a442f4251d8, points: 1, grade: true, solution: false -->
```python
# Test getting minimum element
assert str(heap.get_min()) == "(0, 'D')"

# Test removing minimum element
assert str(heap.remove_min()) == "(0, 'D')"

# Test heap property after removing minimum element
assert str(heap) == "[(1, 'S'), (2, 'E'), (5, 'S'), (4, 'B'), (10, 'A'), (8, 'U'), (9, 'B'), (7, 'A')]"
```

You have to add `max_heap_sort()` function which sorts the current min heap to a max heap before you run the test case cell below.

<!-- grade_id: cell-977ca70f25b93c49, points: 1, grade: true, solution: false -->
```python
# Test sorting the heap to a max heap
heap.max_heap_sort()
print(heap)
assert str(heap) == "[(10, 'A'), (7, 'A'), (9, 'B'), (4, 'B'), (2, 'E'), (8, 'U'), (5, 'S'), (1, 'S')]"
```

Output:
```
[(10, 'A'), (7, 'A'), (9, 'B'), (4, 'B'), (2, 'E'), (8, 'U'), (5, 'S'), (1, 'S')]
```

## **Problem 2: Median of A Numeric Array (3)**
The median is the middle value in an ordered numeric 1D array. If the size of the array is even, there is no middle value and the median is the mean of two middle values.

Given a inordered array, implement a solution which utilize heap data structure to find its median. Note that you cannot use Python built-in heap queue object, try to implement some simple binary heaps to support your solution instead.

<!-- grade_id: cell-c6ed543ce51d1bf3, points: 0, grade: false, solution: true -->
```python
def find_median(nums):
    min_heap = MinHeap()
    max_heap = MinHeap()

    for num in nums:
        if not max_heap._data:
            max_heap.add(Item(-num, None))
            continue
        
        if num < -max_heap.get_min().key:
            max_heap.add(Item(-num, None))
        else:
            min_heap.add(Item(num, None))
        
        if len(max_heap._data) > len(min_heap._data) + 1:
            val = -max_heap.remove_min().key
            min_heap.add(Item(val, None))
        elif len(min_heap._data) > len(max_heap._data):
            val = min_heap.remove_min().key
            max_heap.add(Item(-val, None))
    
    if len(max_heap._data) > len(min_heap._data):
        return -max_heap.get_min().key
    else:
        return (-max_heap.get_min().key + min_heap.get_min().key) / 2

```

<!-- grade_id: cell-9ab03ae69ab89a3e, points: 1.5, grade: true, solution: false -->
```python
assert find_median([0, 2, 5, 7, 9, 4, 3]) == 4
assert find_median([0.5, -1, 3.25, -0.75, 2.5, 0]) == 0.25
```

<!-- grade_id: cell-7dd4aeae9c6ead91, points: 1.5, grade: true, solution: false -->
```python
assert find_median([5]) == 5
```

## **Problem 3: Merge Overlapping Intervals (3)**
Given an array of time intervals where $arr[i] = [start_i, end_i]$, the task is to merge all the overlapping intervals into one and output the result which should have only mutually exclusive intervals.

Examples:
```
Input: arr[] = [[1, 3], [2, 4], [6, 8], [9, 10]]
Output: [[1, 4], [6, 8], [9, 10]]
Explanation: In the given intervals, we have only two overlapping intervals [1, 3] and [2, 4].
             Therefore, we will merge these two and return [[1, 4]], [6, 8], [9, 10]].

Input: arr[] = [[7, 8], [1, 5], [2, 4], [4, 6]]
Output: [[1, 6], [7, 8]]
Explanation: We will merge the overlapping intervals [[1, 5], [2, 4], [4, 6]] into a single interval [1, 6].
```

<!-- grade_id: cell-b16305df828c20a2, points: 0, grade: false, solution: true -->
```python
def merge_intervals(arr):
    arr.sort(key=lambda x: x[0])
    merged = []
    for item in arr:
        if not merged or merged[-1][1] < item[0]:
            merged.append(item)
        else:
            merged[-1][-1] = max(merged[-1][1], item[1])
    
    return merged
```

<!-- grade_id: cell-e45f5daaaa311ef3, points: 1.5, grade: true, solution: false -->
```python
assert merge_intervals([[1, 3], [2, 4], [6, 8], [9, 10]]) == [[1, 4], [6, 8], [9, 10]]
assert merge_intervals([[7, 8], [1, 5], [2, 4], [4, 6]]) == [[1, 6], [7, 8]]
assert merge_intervals([[1, 4], [4, 5]]) == [[1, 5]]
```

<!-- grade_id: cell-13ec7220f2fcad96, points: 1.5, grade: true, solution: false -->
```python
assert merge_intervals([[1,3]]) == [[1,3]]
```

---
# **Hash Map**

# Problem 4: Implement the Hashing class using the Separate Chaining method (4)

<!-- grade_id: cell-6e1766760b517af9, points: 0, grade: false, solution: true -->
```python
class Node:
    def __init__(self, key, value, next=None):
        self.key = key
        self.value = value
        self.next = next

    def __repr__(self):
        return f"({self.key}, '{self.value}')"

class SepChainHash:
    def __init__(self, size):
        self.size = size
        self.data = [None] * self.size

    def hash(self, key):
        return key % self.size

    def add(self, key, value):
        index = self.hash(key)
        current = self.data[index]
        
        prev = None
        while current:
            if current.key == key:
                current.value = value
                return
            prev = current
            current = current.next
        
        new_node = Node(key, value)
        if prev:
            prev.next = new_node
        else:
            self.data[index] = new_node

    def __getitem__(self, key):
        index = self.hash(key)
        current = self.data[index]

        while current:
            if current.key == key:
                return current.value
            current = current.next

        return None

    def delete(self, key):
        index = self.hash(key)
        current = self.data[index]
        previous = None

        while current:
            if current.key == key:
                if previous:
                    previous.next = current.next
                else:
                    self.data[index] = current.next
                return
            previous = current
            current = current.next

    def __repr__(self):
        result = []
        for chain in self.data:
            current = chain
            while current:
                result.append(str(current))
                current = current.next
        return "[" + ", ".join(result) + "]"
```

<!-- grade_id: cell-c9b8958c666c90d4, points: 2, grade: true, solution: false -->
```python
# Test adding items to the hash table
hash_table = SepChainHash(5)
hash_table.add(6, 'a')
hash_table.add(1, 'b')
hash_table.add(12, 'h')
hash_table.add(10, 'r')
hash_table.add(6, 'p')
hash_table.add(4, 's')
hash_table.add(2, 'n')

# Test retrieving item with a non-existing key
assert str(hash_table[9]) == "None", "Test Fail: Key 9 should not exist."

# Test retrieving item with an existing key
assert str(hash_table[4]) == "s", "Test Fail: Key 4 should return 's'."

# Test the hash table state before deleting
print("Before deletion:", hash_table)

# Test deleting an item and check if the chain updates correctly
hash_table.delete(6)
print("After deleting key 6:", hash_table)


```

Output:
```
Before deletion: [(10, 'r'), (6, 'p'), (1, 'b'), (12, 'h'), (2, 'n'), (4, 's')]
After deleting key 6: [(10, 'r'), (1, 'b'), (12, 'h'), (2, 'n'), (4, 's')]
```

<!-- grade_id: cell-794cda9b35b2884d, points: 2, grade: true, solution: false -->
```python
# Create a new hash table
hash_table = SepChainHash(5)

# Test adding multiple items with collision (same hash bucket)
hash_table.add(1, 'apple')
hash_table.add(6, 'banana')
hash_table.add(11, 'cherry')

# Test that the chain at a particular index represents all values (including collisions)
assert str(hash_table) == "[(1, 'apple'), (6, 'banana'), (11, 'cherry')]", "Test Fail: Collision chain incorrect."

# Test getting values by key
assert hash_table[1] == 'apple', "Test Fail: Expected 'apple'."
assert hash_table[6] == 'banana', "Test Fail: Expected 'banana'."
assert hash_table[11] == 'cherry', "Test Fail: Expected 'cherry'."

# Test updating existing key
hash_table.add(1, 'apricot')
assert hash_table[1] == 'apricot', "Test Fail: Expected 'apricot'."

# Test retrieving non-existent key
assert hash_table[15] is None, "Test Fail: Key 15 should return None."

# Test deleting an existing key
hash_table.delete(1)
assert hash_table[1] is None, "Test Fail: Key 1 should have been deleted."

# Test the state of the chain after deletions and additions
hash_table.add(21, 'date')
hash_table.add(16, 'elderberry')
assert str(hash_table) == "[(6, 'banana'), (11, 'cherry'), (21, 'date'), (16, 'elderberry')]", "Test Fail: Chain state incorrect after operations."

# Test deleting a key that does not exist
hash_table.delete(99)
assert str(hash_table) == "[(6, 'banana'), (11, 'cherry'), (21, 'date'), (16, 'elderberry')]", "Test Fail: Deleting non-existent key altered chain."

# Test the complete hash table representation
print("Final hash table state:", hash_table)
# Create a new hash table
hash_table = SepChainHash(5)

# Test adding multiple items with collision (same hash bucket)
hash_table.add(1, 'apple')
hash_table.add(6, 'banana')
hash_table.add(11, 'cherry')

# Test that the chain at a particular index represents all values (including collisions)
assert str(hash_table) == "[(1, 'apple'), (6, 'banana'), (11, 'cherry')]", "Test Fail: Collision chain incorrect."

# Test getting values by key
assert hash_table[1] == 'apple', "Test Fail: Expected 'apple'."
assert hash_table[6] == 'banana', "Test Fail: Expected 'banana'."
assert hash_table[11] == 'cherry', "Test Fail: Expected 'cherry'."

# Test updating existing key
hash_table.add(1, 'apricot')
assert hash_table[1] == 'apricot',   "Test Fail: Expected 'apricot'."

# Test retrieving non-existent key
assert hash_table[15] is None, "Test Fail: Key 15 should return None."

# Test deleting an existing key
hash_table.delete(1)
assert hash_table[1] is None, "Test Fail: Key 1 should have been deleted."

# Test the state of the chain after deletions and additions
hash_table.add(21, 'date')
hash_table.add(16, 'elderberry')
assert str(hash_table) == "[(6, 'banana'), (11, 'cherry'), (21, 'date'), (16, 'elderberry')]", "Test Fail: Chain state incorrect after operations."

# Test deleting a key that does not exist
hash_table.delete(99)
assert str(hash_table) == "[(6, 'banana'), (11, 'cherry'), (21, 'date'), (16, 'elderberry')]", "Test Fail: Deleting non-existent key altered chain."

# Test the complete hash table representation
print("Final hash table state:", hash_table)

```

Output:
```
Final hash table state: [(6, 'banana'), (11, 'cherry'), (21, 'date'), (16, 'elderberry')]
Final hash table state: [(6, 'banana'), (11, 'cherry'), (21, 'date'), (16, 'elderberry')]
```

# Problem 5: Implement a Hashing class with Linear Probing method (4)

<!-- grade_id: cell-84634e5ac3aaaeea, points: 0, grade: false, solution: true -->
```python
class LinProbHash:
    def __init__(self, size):
        self.size = size
        self.data = [None] * self.size

    def hash(self, key):
        return key % self.size

    def add(self, key, value):
        index = self.hash(key)
        original_index = index
        
        while self.data[index] is not None and self.data[index][0] != key:
            index = (index + 1) % self.size
           
            if index == original_index:
                raise ValueError("Hash table is full")
                
        self.data[index] = (key, value)

    def __getitem__(self, key):
        index = self.hash(key)
        original_index = index
        
        while self.data[index] is not None:
            if self.data[index][0] == key:
                return self.data[index][1]
            index = (index + 1) % self.size
            if index == original_index:
                break
                
        return None

    def delete(self, key):
        index = self.hash(key)
        original_index = index
        
        while self.data[index] is not None:
            if self.data[index][0] == key:
                self.data[index] = None
                
                next_index = (index + 1) % self.size
                while self.data[next_index] is not None:
                    key_to_rehash, value_to_rehash = self.data[next_index]
                    self.data[next_index] = None
                    self.add(key_to_rehash, value_to_rehash)
                    next_index = (next_index + 1) % self.size
                    if next_index == original_index:
                        break
                return
                
            index = (index + 1) % self.size
            if index == original_index:
                break

    def __repr__(self):
        return str(self.data)
```

<!-- grade_id: cell-b7152dfd73891dfc, points: 2, grade: true, solution: false -->
```python
# Test adding items to the hash table
lin_prob_hash = LinProbHash(8)
lin_prob_hash.add(6, 's')
lin_prob_hash.add(3, 'd')
lin_prob_hash.add(11, 'a')
lin_prob_hash.add(19, 'p')
lin_prob_hash.add(8, 'd')
lin_prob_hash.add(14, 'q')
lin_prob_hash.add(21, 'j')

# Test the hash table's initial state
print("Initial state:", lin_prob_hash)

# Test deleting items with keys 11 and 3, and check the table state
lin_prob_hash.delete(11)
lin_prob_hash.delete(3)
print("After deleting keys 11 and 3:", lin_prob_hash)

# Test adding more items
lin_prob_hash.add(16, 'l')
lin_prob_hash.add(0, 'o')
lin_prob_hash.add(9, 'a')

# Test deleting another key and check final state
lin_prob_hash.delete(8)
print("Final state after all operations:", lin_prob_hash)

```

Output:
```
Initial state: [(8, 'd'), (21, 'j'), None, (3, 'd'), (11, 'a'), (19, 'p'), (6, 's'), (14, 'q')]
After deleting keys 11 and 3: [(8, 'd'), None, None, (19, 'p'), None, (21, 'j'), (6, 's'), (14, 'q')]
Final state after all operations: [(16, 'l'), (0, 'o'), (9, 'a'), (19, 'p'), None, (21, 'j'), (6, 's'), (14, 'q')]
```

<!-- grade_id: cell-712fe4162f2cfa33, points: 2, grade: true, solution: false -->
```python
# Test initialization and basic operations
hash_table = LinProbHash(10)
hash_table.add(0, 'apple')
hash_table.add(1, 'banana')
hash_table.add(2, 'cherry')
hash_table.add(3, 'date')
hash_table.add(10, 'fig')
assert hash_table[10] == 'fig', "Test Fail: Fig should be at position calculated by wrapping around."

```

# Problem 6: Mystery of the Cryptic Manuscript (2)

You are an explorer who has discovered a mysterious manuscript filled with sequences of letters. This manuscript contains secrets that can only be revealed by understanding the longest sequence of its cryptic text that uses no more than two distinct letters. Your task is to decode the longest such sequence to help unravel the manuscript's mysteries.

Given a string `s` representing a segment of the manuscript, return the length of the longest substring that contains at most two distinct characters. The success of your mission depends on this crucial piece of information!

### Constraints:
- `s` consists of English letters (lowercase and uppercase).


<!-- grade_id: cell-d19cfd7cc710fee1, points: 0, grade: false, solution: true -->
```python
def longest_substring_two_distinct(s: str) -> int:
    if not s:
        return 0
    
    char_count = {}
    max_length = start = 0
    
    for end, char in enumerate(s):
        char_count[char] = char_count.get(char, 0) + 1
        
        while len(char_count) > 2:
            start_char = s[start]
            char_count[start_char] -= 1
            if char_count[start_char] == 0:
                del char_count[start_char]
            start += 1
        
        max_length = max(max_length, end - start + 1)
    
    return max_length
```

<!-- grade_id: cell-966c500c0ef14557, points: 2, grade: true, solution: false -->
```python
assert longest_substring_two_distinct("abracadabra") == 3, "Test 1 Failed"  # "abab"
assert longest_substring_two_distinct("xyyx") == 4, "Test 2 Failed"  # "xyyx"
assert longest_substring_two_distinct("scrolls") == 3, "Test 3 Failed"  # "oll" or "lls"
assert longest_substring_two_distinct("eerie") == 3, "Test 4 Failed"  # "eerie"
assert longest_substring_two_distinct("longitude") == 2, "Test 5 Failed"
assert longest_substring_two_distinct("a") == 1, "Test 6 Failed"  # "a"
assert longest_substring_two_distinct("abcdef") == 2, "Test 7 Failed"  # "ab"
assert longest_substring_two_distinct("balloon") == 4, "Test 8 Failed"  # "lloo"
```

```python
import heapq

# =======================
# Tạo danh sách ban đầu
# =======================
data = [5, 2, 8, 1, 9, 3]

# =========================================
# 1. heapq.heapify(list): Biến list thành Min Heap
# =========================================
heapq.heapify(data)
print("1. Heap sau heapify:", data)  # Phần tử nhỏ nhất ở đầu

# =========================================
# 2. heapq.heappush(heap, item): Thêm phần tử mới vào heap
# =========================================
heapq.heappush(data, 0)
print("2. Sau heappush(0):", data)

# =========================================
# 3. heapq.heappop(heap): Lấy ra và xoá phần tử nhỏ nhất
# =========================================
smallest = heapq.heappop(data)
print("3. Phần tử nhỏ nhất bị loại:", smallest)
print("   Heap còn lại:", data)

# ==================================================
# 4. heapq.heappushpop(heap, item):
#     → Thêm item vào heap rồi lấy ra phần tử nhỏ nhất
#     → Hiệu quả hơn khi dùng kết hợp heappush + heappop
# ==================================================
result = heapq.heappushpop(data, 4)
print("4. Kết quả heappushpop(4):", result)
print("   Heap sau thao tác:", data)

# ==================================================
# 5. heapq.heapreplace(heap, item):
#     → Loại phần tử nhỏ nhất trước, rồi thêm item mới vào heap
#     → Hiệu quả hơn dùng pop + push
# ==================================================
result = heapq.heapreplace(data, 10)
print("5. Kết quả heapreplace(10):", result)
print("   Heap sau thao tác:", data)

# =========================================
# 6. heapq.nlargest(n, iterable):
#     → Lấy n phần tử lớn nhất từ danh sách
# =========================================
largest = heapq.nlargest(3, [1, 4, 9, 2, 7, 5])
print("6. 3 phần tử lớn nhất:", largest)

# =========================================
# 7. heapq.nsmallest(n, iterable):
#     → Lấy n phần tử nhỏ nhất từ danh sách
# =========================================
smallest = heapq.nsmallest(3, [1, 4, 9, 2, 7, 5])
print("7. 3 phần tử nhỏ nhất:", smallest)

# =========================================
# 8. Mô phỏng Max Heap bằng cách đảo dấu (âm)
# =========================================
max_heap = []
heapq.heappush(max_heap, -5)
heapq.heappush(max_heap, -1)
heapq.heappush(max_heap, -8)
heapq.heappush(max_heap, -3)

# Lấy phần tử lớn nhất (phải đảo dấu lại)
largest = -heapq.heappop(max_heap)
print("8. Phần tử lớn nhất (mô phỏng Max Heap):", largest)
```