Make sure you :
- Fill in your name
- Fill any place that says `YOUR CODE HERE` (you will be graded based on this code, any extra cells you code will not be graded)
- Delete  `NotImplementedError()` exception once you have completed your implementation

```python
NAME = "NGUYEN THI THU TRANG"
```

# Problem 1: Recursion Fundamentals (4)

Give the appropriate answers for the questions below, **you only need 1 sentence for each question**.

1. What is Recursion? How Recursion Works.
2. How do you determine the base case in a recursive function?
3. Compare Recursion vs. Iteration?
4. What is the advantage of recursion over the iterative (loop) approach?


<!-- grade_id: cell-0bd1a86e41779b9a, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

1. Recursion is a programming technique where a function calls itself to solve a smaller version of the original problem until it reaches a base case.
2. The base case is determined by identifying the simplest input for which the function can return a result without making further recursive calls.
3. Recursion uses function calls and stack memory to solve problems, while iteration uses loops and modifies variables directly without additional function calls.
4. The advantage of recursion is that it simplifies complex problems, especially those involving tree structures or backtracking, making the code more readable and easier to understand.

# Problem 2: Counting Digits (1)

Write a recursive function `count_digits(n)` that takes a positive integer n and returns the number of digits in n.

You are not allowed to use loops, string conversions, or in-built functions to calculate the digit count.


<!-- grade_id: cell-8c0d4facb2fc900f, points: 0, grade: false, solution: true -->
```python
def count_digits(n):
    if n < 10:
        return 1
    return 1 + count_digits(n // 10)
```

<!-- grade_id: cell-f4997c6dd443359c, points: 0.5, grade: true, solution: false -->
```python
assert count_digits(12345) == 5
assert count_digits(9) == 1
```

<!-- grade_id: cell-b704a8df16c6e979, points: 0.5, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```

# Problem 3: Palindrome Check (2)

A palindrome is a word, phrase, number, or other sequence of characters that reads the same forward and backward.

Example:
- **Radar** is a palindrome
- **hello** is not a palindrome

Write a recursive function to determine a palindrome.

Hint: To determine a palindrome using recursion, the outer characters of a word are compared. If they match, the inner substring is checked recursively.


<!-- grade_id: cell-12a7aa3c42e418a5, points: 0, grade: false, solution: true -->
```python
def is_palindrome(word) -> bool:
   if len(word) <= 1:
      return True
   if word[0].lower() != word[-1].lower():
      return False
   return is_palindrome(word[1:-1])
```

<!-- grade_id: cell-ee59b83efaf83891, points: 0.5, grade: true, solution: false -->
```python
assert is_palindrome("Radar") == True
assert is_palindrome("hello") == False
assert is_palindrome("123 isi 321") == True
```

<!-- grade_id: cell-12c1d111ba288c27, points: 0.75, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```

<!-- grade_id: cell-dbff95ebfead83e7, points: 0.75, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```

# Problem 4: Catalan Numbers (2)

The Catalan numbers form a sequence of natural numbers that occur in various counting problems, often involving recursive structures. The recursive formula for the $n$-th Catalan number is

$$C_n = \sum_{i=0}^{n-1} C_i ⋅ C_{n-1-i}$$

where the base case is $C_0 = 1$

Write a recursive function to determine the $n$-th Catalan number.

<!-- grade_id: cell-fb2dbceb03b04dc8, points: 0, grade: false, solution: true -->
```python
def catalan_number_memoized(n, memo={}):
    if n == 0:
        return 1
    return sum(catalan_number_memoized(i) * catalan_number_memoized(n - 1 - i) for i in range(n))

```

<!-- grade_id: cell-c410465035467a04, points: 0.25, grade: true, solution: false -->
```python
assert catalan_number_memoized(0) == 1, "Test case failed: C(0) should be 1"
```

<!-- grade_id: cell-f6ce5254d51797a4, points: 0.25, grade: true, solution: false -->
```python
# Small numbers
assert catalan_number_memoized(2) == 2, "Test case failed: C(2) should be 2"
# Moderate numbers
assert catalan_number_memoized(7) == 429, "Test case failed: C(7) should be 429"
```

<!-- grade_id: cell-70ef06f43f19996d, points: 0.75, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```

<!-- grade_id: cell-e1588820262882b2, points: 0.75, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL

```

# Problem 5: Generating Valid Parentheses (1)

Write a recursive function `generate_parentheses(n)` that takes an integer n and returns a list of all valid combinations of n pairs of parentheses.

You must ensure that only valid parentheses combinations are generated (e.g., no "())(" or "(())())").

Example:
```
Input: n = 2
Output: ['(())', '()()']
Note:['()()', '(())'] also a valid answer, do not worry about the order.
```

<!-- grade_id: cell-7c80cefea260b3a2, points: 0, grade: false, solution: true -->
```python
def generate_parentheses(n):
    stack = []
    res = []

    def backtrack(open, closed):
        if open == closed == n:
            res.append("".join(stack))
            return
        
        if open < n:
            stack.append("(")
            backtrack(open + 1, closed)
            stack.pop()
        
        if closed < open:
            stack.append(")")
            backtrack(open, closed + 1)
            stack.pop()


    backtrack(0, 0)
    return res

   
```

<!-- grade_id: cell-c7c67a1dd15a8446, points: 0.5, grade: true, solution: false -->
```python
assert sorted(generate_parentheses(2)) == ['(())', '()()']
assert sorted(generate_parentheses(3)) == ['((()))', '(()())', '(())()', '()(())', '()()()']
```

<!-- grade_id: cell-0ded4a0379134295, points: 0.25, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```

<!-- grade_id: cell-9f52a5f240c85628, points: 0.25, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```