Make sure you :
- Fill in your name
- Fill any place that says `YOUR CODE HERE` (you will be graded based on this code, any extra cells you code will not be graded)
- Delete  `NotImplementedError()` exception once you have completed your implementation

```python
NAME = "NGUYEN THI THU TRANG"
```

# Problem 1: Array Fundamentals (4)

Give the appropriate answers for the questions below, **you only need 1-2 sentences for each question**.

1. What is an Array in Programming?
2. Mention some advantages and disadvantages of Arrays.
3. How is data stored in an array?
4. What happens when an array exceeds its allocated size?


ANSWER HERE (click here twice to answer)

1.	An array in programming is a data structure that stores multiple elements of the same type in a contiguous memory block, allowing efficient access via indexing.
2.	Advantages: Arrays provide fast access to elements and efficient iteration.   
    Disadvantages: They have a fixed size in many languages and require shifting elements for insertions or deletions.
3.	Data in an array is stored in contiguous memory locations, enabling fast element access using an index.
4.	In many languages, exceeding an array’s allocated size causes memory errors, but in Python, lists dynamically resize by allocating more memory and copying elements.

# Problem 2: Insertion Sort Algorithm (2)
Implement Insertion Sort Algorithm to sort a list in ascending order.

<!-- grade_id: cell-33d48ef303de6d80, points: 0, grade: false, solution: true -->
```python
def insertion_sort(arr):
    for i in range(1, len(arr)):
        key = arr[i]
        j = i - 1

        while j >= 0 and key < arr[j]:
            arr[j + 1] = arr[j]
            j -= 1
        arr[j + 1] = key
    
    return arr

```

<!-- grade_id: cell-3ed5ec8fbb6a227a, points: 1, grade: true, solution: false -->
```python
assert insertion_sort([4, 2, 1, 3]) == [1, 2, 3, 4]
assert insertion_sort([-1, 5, -3, 0]) == [-3, -1, 0, 5]
assert insertion_sort([1, 2, 1, 2]) == [1, 1, 2, 2]
```

<!-- grade_id: cell-000fa2215890218e, points: 1, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```

# Problem 3: Remove duplicates from Sorted Array (2)

Write a function `removeDuplicates(arr)` that takes a sorted array arr[] (may contain duplicate elements) and return a sorted array with duplicate elements removed.

**Note:**
- Try to run in O(n) Time and O(1) Space.

**Examples:**

Input: arr = [1, 2, 2, 3, 4, 4, 4, 5, 5]\
Output: [1, 2, 3, 4, 5]

<!-- grade_id: cell-7b82c5a765b4a113, points: 0, grade: false, solution: true -->
```python
def removeDuplicates(arr):
    if not arr:
            return 0

    write_index = 1
    for i in range(1, len(arr)):
        if arr[i] != arr[i - 1]:
            arr[write_index] = arr[i]
            write_index += 1

    return arr[:write_index]
```

<!-- grade_id: cell-f87088353c4efd7d, points: 1, grade: true, solution: false -->
```python
assert removeDuplicates([1,2,2,3,4,4,4,5,5]) == [1,2,3,4,5]
assert removeDuplicates([1,2,3,4,5]) == [1,2,3,4,5]
```

<!-- grade_id: cell-c511d9e2b620db7a, points: 1, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```

# Problem 4: Zero In Total (2)

Given an integer array nums, return all the triplets `[nums[i], nums[j], nums[k]]` such that **i != j, i != k, and j != k**, and **nums[i] + nums[j] + nums[k] == 0.**

Notice that the solution set must not contain duplicate triplets.

Example 1:
```
Input: nums = [-1,0,1,2,-1,-4]
Output: [[-1,-1,2],[-1,0,1]]
```
Explanation:

1. nums[0] + nums[1] + nums[2] = (-1) + 0 + 1 = 0.
2. nums[1] + nums[2] + nums[4] = 0 + 1 + (-1) = 0.
3. nums[0] + nums[3] + nums[4] = (-1) + 2 + (-1) = 0.
- The distinct triplets are [-1,0,1] and [-1,-1,2].
- Notice that the order of the output and the order of the triplets does not matter.

Constraints:
- 3 <= nums.length <= 3000
- -105 <= nums[i] <= 105


<!-- grade_id: cell-9d1df92200c81020, points: 0, grade: false, solution: true -->
```python
def zeroTotal(nums: list[int]) -> list[list[int]]:
    nums.sort()
    ans = []
    for i, num in enumerate(nums):
        if i > 0 and num == nums[i - 1]:
            continue

        left = i + 1
        right = len(nums) - 1
        while left < right:
            total = num + nums[left] + nums[right]
            if  total > 0:
                right -= 1
            elif total < 0:
                left += 1
            else:
                ans.append([num, nums[left], nums[right]])
                left += 1
                while nums[left] == nums[left - 1] and left < right:
                    left += 1

    return ans

```

```python
def normalize(ans):
  return sorted([sorted(triplet) for triplet in ans])
```

```python
assert normalize(zeroTotal([-1,0,1,2,-1,-4])) == [[-1, -1, 2], [-1, 0, 1]]
```

<!-- grade_id: cell-27833fceb8b7cb94, points: 1, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL

```

<!-- grade_id: cell-4fd7100c251fe0c0, points: 1, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```