Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you :
- Fill in your name
- Fill any place that says `YOUR CODE HERE` (you will be graded based on this code, any extra cells you code will not be graded)
- Delete  `NotImplementedError()` exception once you have completed your implementation

```python
NAME = "NGUYEN THI THU TRANG"
```

# Problem 1: Stack Fundamentals (4)

Give the appropriate answers for the questions below, **you only need 1-2 sentences for each question**.

1. What is a Stack?
2. What are the primary operations performed on a Stack and their time complexities?
3. Why is Stack considered a Recursive data structure?
4. What are Practical Applications of Stack?

ANSWER HERE (click here twice to answer)
1. A Stack is a linear data structure that follows the Last In, First Out (LIFO) principle.
2. The main operations are push (O(1)), pop (O(1)), top (O(1)), and is_empty (O(1)).
3. Stack is considered recursive because function calls in Python use an implicit call stack.
4. Stacks are used in recursion, expression evaluation, undo/redo functionality, and backtracking.

# Problem 2: Stack from Scratch (3)

## Part 1: Stack Implementation (2)
Implement class `Stack` you learned in the class again with below structure:
- `__init__(self)`: which initialize an empty stack
- `__len__(self)`: return the number of elements in the stack
- `__str__(self)`: Return string presentation of the stack
- `is_empty(self)`: Return True if the stack is empty and vice versa
- `push(self, a)`: Add element a to the top of the stack
- `top(self)`: Return (but do not remove) the element at the top of the stack
- `pop(self)`: Remove and return the element from the top of the stack

<!-- grade_id: cell-4693bdab113c7887, points: 0, grade: false, solution: true -->
```python
class Stack:
    def __init__(self):
        self.data = []
    
    def __len__(self):
        return len(self.data)
    
    def __str__(self):
        return str(self.data)
    
    def is_empty(self):
        return len(self.data) == 0
    
    def push(self, a):
        self.data.append(a)

    def top(self):
        return self.data[-1]
    
    def pop(self):
        return self.data.pop()
    
    
```

```python
A = Stack()
A.push(4)
A.push(0)
A.push(1)
A.push(5)
A.push(9)
A.push(8)
A.push(2)
```

<!-- grade_id: cell-456377f109dc13d5, points: 2, grade: true, solution: false -->
```python
assert A.is_empty() == False
assert str(A) == '[4, 0, 1, 5, 9, 8, 2]'
assert len(A) == 7
assert A.top() == 2
assert A.pop() == 2
```

## Part 2: Moving Elements (1)
Build a function to move all elements from stack A to stack B without changing the order

<!-- grade_id: cell-1b0790b18986aa37, points: 0, grade: false, solution: true -->
```python
def moving_elements(stackA: Stack, stackB: Stack) -> Stack:
    temp = Stack()

    while not stackA.is_empty():
        temp.push(stackA.pop())

    while not temp.is_empty():
        stackB.push(temp.pop())

    return stackB
```

```python
B = Stack()
moving_elements(A, B)
```

Output:
```
<__main__.Stack at 0x1060fbef0>
```

<!-- grade_id: cell-3039332d546c9e2f, points: 1, grade: true, solution: false -->
```python
assert str(B) == '[4, 0, 1, 5, 9, 8]'
```

# Problem 3: Evaluation of Postfix Expression (2)
Given a postfix expression, the task is to evaluate the postfix expression. A Postfix expression is of the form “a b operator” (“a b +”) i.e., a pair of operands is followed by an operator.

**Examples:**
```
Input: arr = [“2”, “3”, “1”, “*”, “+”, “9”, “-“]
Output: -4
Explanation: If the expression is converted into an infix expression, it will be 2 + (3 * 1) – 9 = 5 – 9 = -4.
```
```
Input: arr = [“100”, “200”, “+”, “2”, “/”, “5”, “*”, “7”, “+”]
Output: 757
Explanation: If the expression is converted into an infix expression, it will be ((100 + 200) / 2) * 5 + 7  = 150 * 5 + 7 = 757.
```
Complete the function below (**Note**: Handle all cases, raise type error if postfix expression is invalid)

<!-- grade_id: cell-5e00d6b4b9567b23, points: 0, grade: false, solution: true -->
```python
def evaluate_postfix(arr):
    if not arr:
        return TypeError("Invalid")
    stack = Stack()
    operators = ['+', '-', '*', '/']

    for element in arr:
        if element.isdigit():
            stack.push(int(element))
        else:
            if element in operators:
                val2 = stack.pop()
                val1 = stack.pop()

                if element == '-':
                    stack.push(val1 - val2)
                elif element == '+':
                    stack.push(val1 + val2)
                elif element == '*':
                    stack.push(val1 * val2)
                elif element == '/':
                    stack.push(val1 / val2)
            else:
                return TypeError("Invalid")
            
    return stack.pop()
```

<!-- grade_id: cell-7d275ce22a65b581, points: 1.5, grade: true, solution: false -->
```python
assert evaluate_postfix(["2", "3", "1", "*", "+", "9", "-"]) == -4
assert evaluate_postfix(["100", "200", "+", "2", "/", "5", "*", "7", "+"]) == 757
```

<!-- grade_id: cell-274041bc2364eb2a, points: 0.5, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```

# Problem 4: Valid Parentheses(1)

Given a string s containing just the characters '(', ')', '{', '}', '[' and ']', determine if the input string is valid.

An input string is valid if:

1. Open brackets must be closed by the same type of brackets.
2. Open brackets must be closed in the correct order.
3. Every close bracket has a corresponding open bracket of the same type.

Example 1:
```
Input: s = "()[]{}"
Output: True

```
Example 2:
```
Input: s = "(]"
Output: False
```

```python
def isValid(s):
    """
    :type s: str
    :rtype: bool
    """
    pairs = {')': '(', '}': '{', ']': '['}
    stack = Stack()
    
    for char in s:
        if char in '[({':
            stack.push(char)
        elif char in '])}':
            if stack.is_empty():
                return False
            if stack.pop() != pairs[char]:
                return False
    
    return stack.is_empty()
        
    
```

<!-- grade_id: cell-d00bd588e5d8b275, points: 0.5, grade: true, solution: false -->
```python
assert isValid("()[]{}") == True
assert isValid('(]') == False
```

<!-- grade_id: cell-7462dbe3df0b291c, points: 0.25, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```

<!-- grade_id: cell-b8e7c506188d49a6, points: 0.25, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```