Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you :
- Fill in your name
- Fill any place that says `YOUR CODE HERE` (you will be graded based on this code, any extra cells you code will not be graded)
- Delete  `NotImplementedError()` exception once you have completed your implementation

```python
NAME = "NGUYEN THI THU TRANG"
```

# Problem 1: Queue Fundamentals (3)

Give the appropriate answers for the questions below, **you only need 1-2 sentences for each question**.

1. What is a Queue? Name some Types of Queues.
2. What are the primary operations performed on a Queue and their time complexities?
3. What are the Practical Applications of Queue?
4. When should we use a Stack or a Queue instead of Arrays/Lists?


ANSWER HERE (click here twice to answer)
1. A Queue is a linear data structure that follows the First In, First Out (FIFO) principle. Types of Queues include Simple Queue, Circular Queue, Priority Queue, and Deque (Double-Ended Queue).
2. The main operations are enqueue (O(1)), dequeue (O(1)), peek (O(1)), and is_empty (O(1)).
3. Queues are used in task scheduling (CPU scheduling, printer queue), BFS traversal, and handling real-time requests in servers.
4. Use a Stack when Last In, First Out (LIFO) behavior is needed (e.g., undo operations, recursion handling). Use a Queue when First In, First Out (FIFO) behavior is needed (e.g., task processing, order management).

# Problem 2: Queue from Scratch (4)

## Part 1: Queue Implementation (2)
Implement class `MyQueue` you learned in the class again with below structure:
- `__init__(self)`: which initialize an empty queue
- `__len__(self)`: return the number of elements in the queue
- `__str__(self)`: Return string presentation of the queue
- `empty(self)`: Return True if the queue is empty and vice versa
- `push(self, a)`: Pushes element x to the back of the queue
- `peek(self)`:  Returns the element at the front of the queue
- `pop(self)`: Removes the element from the front of the queue and returns it

Use only a list to implement the queue. Do not use any other data structure.

<!-- grade_id: cell-95455757fe0603f1, points: 0, grade: false, solution: true -->
```python
class MyQueue:
    def __init__(self):
        self.data = []

    def __len__(self):
        return len(self.data)
    
    def __str__(self):
        return str(self.data)
    
    def is_empty(self):
        return len(self.data) == 0
    
    def push(self, a):
        self.data.append(a)

    def peek(self):
        return self.data[0]
    
    def pop(self):
        return self.data.pop(0)
    
    def sum(self):
        return sum(self.data)
```

<!-- grade_id: cell-17289684257e01ca, points: 1, grade: true, solution: false -->
```python
Q = MyQueue()
for i in [4, 0, 1, 5, 9, 8, 2]:
    Q.push(i)
assert Q.is_empty() == False
assert str(Q) == '[4, 0, 1, 5, 9, 8, 2]'
assert len(Q) == 7
```

<!-- grade_id: cell-6bd68c7531df56b4, points: 1, grade: true, solution: false -->
```python
assert Q.peek() == 4
assert Q.pop() == 4
```

## Part 2: Stack from Queues (2)
Implement a stack with LIFO behavior using two queues. The stack should include these functions:

- `void push(int x)`: Adds x to the stack's top.
- `int pop()`: Removes and returns the top element.
- `int top()`: Retrieves the top element.
- `boolean empty()`: Checks if the stack is empty, returning true or false.

Follow these guidelines:

- Utilize only standard queue operations: enqueue (push to back), dequeue (peek/pop from front), size, and check if empty.
- If your programming language doesn't support queues, simulate them using lists or deques, adhering to queue's standard behaviors.

<!-- grade_id: cell-95bf6f7e7e21f3e4, points: 0, grade: false, solution: true -->
```python
class MyStack:
    def __init__(self):
        self.q1 = MyQueue()
        self.q2 = MyQueue()
    
    def push(self, x):
        self.q2.push(x)
        while not self.q1.is_empty():
            self.q2.push(self.q1.pop())
        
        self.q1, self.q2 = self.q2, self.q1
    
    def pop(self):
        if not self.q1.is_empty():
            return self.q1.pop()
    
    def top(self):
        if not self.q1.is_empty():
            return self.q1.peek()

    def is_empty(self):
        return self.q1.is_empty()
```

<!-- grade_id: cell-30e61bfffd5d6e6f, points: 1, grade: true, solution: false -->
```python
stack = MyStack()
stack.push(1)
stack.push(2)
assert stack.top() == 2  # Verifies the top element is 2
assert stack.pop() == 2  # Verifies popping returns 2 and removes it
```

<!-- grade_id: cell-b5c538c2e7f74c22, points: 0.5, grade: true, solution: false -->
```python
stack.push(3)
assert stack.pop() == 3
assert not stack.is_empty()
```

<!-- grade_id: cell-0e975e7ca8220266, points: 0.5, grade: true, solution: false -->
```python
stack.push(4)
stack.push(5)
assert stack.pop() == 5
assert stack.top() == 4
assert stack.pop() == 4
assert stack.pop() == 1
assert stack.is_empty()
```

# Problem 3: Reversing a queue using recursion (2)
Given a queue, write a recursive function to reverse it (Using your Queue Implementation).

**Example**
```
Input : Q = [5, 24, 9, 6, 8, 4, 1, 8, 3, 6]
Output : Q = [6, 3, 8, 1, 4, 8, 6, 9, 24, 5]
```

<!-- grade_id: cell-1a9af008ed06cf6f, points: 0, grade: false, solution: true -->
```python
def reverse_queue(q: MyQueue):
    if q.is_empty():
        return q
    else:
        temp = q.peek()
        q.pop()
        q = reverse_queue(q)
        q.push(temp)

        return q
    

```

```python
q = MyQueue()
for i in [5, 24, 9, 6, 8, 4, 1, 8, 3, 6]:
    q.push(i)
```

<!-- grade_id: cell-de0d83b5059d0785, points: 1, grade: true, solution: false -->
```python
reverse_queue(q)
assert str(q) == "[6, 3, 8, 1, 4, 8, 6, 9, 24, 5]"
```

<!-- grade_id: cell-a50eba0dab87ffa2, points: 0.5, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```

<!-- grade_id: cell-1c16d9aee339b6d9, points: 0.5, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```

# Problem 3: Classmates Teaching Each Other (1)

On day 1, student Trong mastered a new concept in Machine Learning. Trong is commited to education and thus after `delay` days, he will teach the concept to another student. However, Trong will forget the concept after `memory` days.

You are given an integer delay, which means that each student will teach that concept with another student every day, starting from delay days after he has learned the new concept. You are also given an integer forget, which means that after forget day of learning it, the student will forget about it. A person cannot share what they learned on the same day they forgot it, or on any day afterwards. Return the number of student who knows the concept on the `day`th day.

```python
def master_machine_learning(day: int, delay: int, memory: int) -> int:
    d = MyQueue() #ready to tell
    f = MyQueue() #ready to forget

    d.push(1) 
    f.push(1)
    
    sharing = 0 
    
    for _ in range(day - 1):
        if d.__len__() >= delay:
            sharing += d.pop()
        if f.__len__() >= memory:
            sharing -= f.pop()
            
        d.push(sharing)  
        f.push(sharing)  

    return f.sum()

```

<!-- grade_id: cell-2eae545c94006764, points: 0.5, grade: true, solution: false -->
```python
assert master_machine_learning(6, 2, 4) == 5
```

<!-- grade_id: cell-96404abe2dc8c07a, points: 0.25, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```

<!-- grade_id: cell-68d18ca5d138f002, points: 0.25, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```