# **Problem 1: Tree Fundamentals (3)**

Give the appropriate answers for the questions below, **you only need 1-2 sentences for each question**.
1. What is a Tree Data Structure? What is a Binary Tree? What are the components of a tree?
2. Explain Height and Depths in the context of a Tree.
3. What is the difference between a Tree and a Graph?
4. What is the worse-case time complexity for a binary search tree for searching? Suggest a solution.

<!-- grade_id: cell-d236c1b60b66ad72, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE
1.	A Tree Data Structure is a hierarchical structure consisting of nodes connected by edges, with one node designated as the root. A Binary Tree is a tree in which each node has at most two children. The components of a tree are the root, nodes, edges, and leaves.
2.	In the context of a Tree, the height is the length of the longest path from a node to a leaf, while the depth is the number of edges from the root to a node.
3.	The main difference between a Tree and a Graph is that a Tree is a connected acyclic graph with a hierarchical structure, while a Graph can have cycles and does not necessarily have a hierarchical structure.
4.	The worst-case time complexity for searching in a binary search tree is O(n) when the tree is unbalanced. A solution is to use a balanced binary search tree or height-balanced tree.

# **Problem 2: BinaryTree class** (2)
## Part a: Implementation

Implement a `Node` class and a `BinaryTree` class with the following specifications:

### Node Class

Each node should contain:
- Its value
- Parent
- Left child
- Right child

### BinaryTree Class

The `BinaryTree` class should include common methods such as:
- `add_left(node, value)`: Adds a left child to a specified node. Raises an Exception if the node already has a left child.
- `add_right(node, value)`: Adds a right child to a specified node. Raises an Exception if the node already has a right child.
- `is_root(node)`: Checks if a node is the root of the tree.
- `is_leaf(node)`: Checks if a node is a leaf (no children).
- `check_ancestor(node1, node2)`: Checks if `node1` is an ancestor of `node2`.

Additionally, implement:
- `height(node)`: Method to find the height of the tree from a given node.
- `depth(node)`: Method to find the depth of a given node in the tree.

For printing:
- Implement `__str__` or `__repr__` methods to print out all elements in the tree in a friendly format.

**Note:**
- In the `add_left` and `add_right` methods, you should raise an appropriate Exception if the node already has a left or right child respectively.
- **Optional:** Bonus points for printing out elements in a user-friendly format.

## Part b: Verification

Check your implementation by performing the following tasks:
- Create a tree with this structure:
```
     1
   /   \
  2     3
 / \    /
4  5   6
  / \ /
  7 8 9

## Part c: Leaf node

Implement function `sum_right_leaves` to class BinTree to return the sum of all the **leaf** nodes that are the right child of their parent in a given binary tree.


#### Input:
```
    1
   / \
  2   3
 / \   \
4   5   6
 \     / \
  7   8   9
```

#### Output:
```
5 + 7 + 9 = 21
```

<!-- grade_id: cell-293da49e24aa0e52, points: 0, grade: false, solution: true -->
```python
class Node:
    def __init__(self, value, parent=None, left_child=None, right_child=None):
        self.value = value
        self.parent = parent
        self.left_child = left_child
        self.right_child = right_child
```

<!-- grade_id: cell-6c2bcdfa1d41039c, points: 0, grade: false, solution: true -->
```python
class BinaryTree:
    def __init__(self, root=None):
        self.root = root

    def add_left(self, p_node: Node, c_node: Node):
        if p_node.left_child is None:
            p_node.left_child = c_node
            c_node.parent = p_node
        else:
            raise Exception("This node already has a left child.")

    def add_right(self, p_node: Node, c_node: Node):
        if p_node.right_child is None:
            p_node.right_child = c_node
            c_node.parent = p_node
        else:
            raise Exception("This node already has a left child.")

    def is_root(self, node: Node):
        return self.root == node

    def is_leaf(self, node: Node):
        return node is not None and node.left_child is None and node.right_child is None

    def check_ancestor(self, node1: Node, node2: Node):
        if node1.left_child == node2 or node1.right_child == node2:
            return True
        return False

    def depth(self, node: Node):
        if node is None:
            return 
        
        depth = 0

        while node != self.root and node:
            node = node.parent
            depth += 1
        
        return depth

    def height(self, node: Node):
        if node is None:
            return -1
        
        left_height = self.height(node.left_child)
        right_height = self.height(node.right_child)

        return max(left_height, right_height) + 1

    def sum_right_leaves(self, node: Node):
        if node is None:
            return 0
        
        sum_val = 0

        if node.right_child and self.is_leaf(node.right_child):
            sum_val += node.right_child.value
        
        return sum_val + self.sum_right_leaves(node.left_child) + self.sum_right_leaves(node.right_child)
    
    def __str__(self):
        return self._print_tree(self.root, 0)

    def _print_tree(self, node, depth):
        if node is None:
            return ""

        right_part = self._print_tree(node.right_child, depth + 1)
        node_part = "   " * depth + str(node.value) + "\n"
        left_part = self._print_tree(node.left_child, depth + 1)
        return right_part + node_part + left_part

```

```python
# Verification - Creating the given tree structure
n1 = Node(1)
n2 = Node(2)
n3 = Node(3)
n4 = Node(4)
n5 = Node(5)
n6 = Node(6)
n7 = Node(7)
n8 = Node(8)
n9 = Node(9)

tree = BinaryTree(n1)
tree.add_left(n1, n2)
tree.add_right(n1, n3)
tree.add_left(n2, n4)
tree.add_right(n2, n5)
tree.add_left(n5, n7)
tree.add_right(n5, n8)
tree.add_left(n3, n6)
tree.add_left(n6, n9)
print(tree)
```

Output:
```
3
      6
         9
1
         8
      5
         7
   2
      4
```

<!-- grade_id: cell-8d09fcfae2dd762f, points: 1, grade: true, solution: false -->
```python
# Check part b
assert tree.depth(n5) == 2
assert tree.height(n5) == 1
assert tree.depth(n9) == 3
assert tree.height(n9) == 0
```

<!-- grade_id: cell-d3e9ea077c681821, points: 1, grade: true, solution: false -->
```python
n1 = Node(1)
n2 = Node(2)
n3 = Node(3)
n4 = Node(4)
n5 = Node(5)
n6 = Node(6)


tree = BinaryTree(n1)
tree.add_left(n1, n2)
tree.add_right(n1, n3)
tree.add_left(n2, n4)
tree.add_right(n2, n5)
tree.add_left(n3, n6)
print(tree)

assert tree.sum_right_leaves(n1) == 5
```

Output:
```
3
      6
1
      5
   2
      4
```

# **Problem 3: Binary Search Tree** (1)

---



Implement and work with a Binary Search Tree (BST), a node-based binary tree data structure with specific properties.

## Properties of a Binary Search Tree

- The left subtree of a node contains only nodes with keys lesser than the node’s key.
- The right subtree of a node contains only nodes with keys greater than the node’s key.
- Both the left and right subtrees must also be binary search trees.
- There must be no duplicate nodes in the tree.

## Tasks

### a. Implementation

1. **Insert Function**: Build a function `insert` to add a new value into the BST. Ensure that the tree maintains its BST properties after the insertion.
   
2. **Iterative Search Function**: Implement a function `iterative_search` to check if a given number is in the BST or not. This function should perform the search without using recursion.

### b. Verification

1. **Create a BST** using the `insert` function. You may specify the series of values to be inserted to form the BST or leave it to the implementer's discretion.

2. **Check for Presence**: Using the `iterative_search` function, check if the values 0, 3, 4, and 12 are present in the BST.

## Notes

- Ensure your `insert` function correctly maintains the BST properties, especially when dealing with edge cases like inserting into an empty tree or inserting values that require traversing to the deeper levels of the tree.
- Your `search` function should efficiently navigate through the tree, leveraging BST properties to minimize the number of comparisons needed to find a value.

<!-- grade_id: cell-d2989d23663e2848, points: 0, grade: false, solution: true -->
```python
class BST(BinaryTree):
    def __init__(self, root = None):
        super().__init__(root)

    def insert(self, num):
        new_node = Node(num)

        if self.root is None:
            self.root = new_node
        
        curr = self.root
        while True:
            if num < curr.value:
                if curr.left_child is None:
                    curr.left_child = new_node
                    new_node.parent = curr
                    break
                curr = curr.left_child
            if num > curr.value:
                if curr.right_child is None:
                    curr.right_child = new_node
                    new_node.parent = curr
                    break
                curr = curr.right_child

    def search(self, value):
        if not self.root:
            return False
        
        curr = self.root
        while curr is not None:
            if curr.value == value:
                return True
            elif curr.value > value:
                curr = curr.left_child
            else:
                curr = curr.right_child

        return False
```

```python
a = Node(8)

bin_search_tree = BST(a)
for num in [3, 1, 6, 4, 7, 10, 14, 13]:
    bin_search_tree.insert(num)

print(bin_search_tree)
```

Output:
```
14
         13
   10
8
         7
      6
         4
   3
      1
```

<!-- grade_id: cell-61783c673c286c76, points: 0.5, grade: true, solution: false -->
```python
assert bin_search_tree.search(0) == False
assert bin_search_tree.search(3) == True
```

<!-- grade_id: cell-135531c778d856a1, points: 0.5, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```

# **Prbolem 4: Twin-tree** (2)

Given the roots of two binary trees p and q, write a function to check if they are the same or not.

Two binary trees are considered the same if they are structurally identical, and the nodes have the same value.

<!-- grade_id: cell-048875c3838d3599, points: 0, grade: false, solution: true -->
```python
def isSameTree(p: BinaryTree, q: BinaryTree):
    """
    :type p: Tree root
    :type q: Tree root
    :rtype: bool

    """
    if not p and not q:
        return True
    if not p or not q or p.value != q.value:
        return False
    
    return isSameTree(p.left_child, q.left_child) and isSameTree(p.right_child, q.right_child)

```

<!-- grade_id: cell-efe08388d7c595ff, points: 0.5, grade: true, solution: false -->
```python
p1 = Node(1)
p2 = Node(2)
p3 = Node(3)
p4 = Node(4)
p5 = Node(5)

tree_P = BinaryTree(p1)
tree_P.add_left(p1, p2)
tree_P.add_right(p1, p3)
tree_P.add_left(p2, p4)
tree_P.add_right(p2, p5)

q1 = Node(1)
q2 = Node(2)
q3 = Node(3)
q4 = Node(4)
q5 = Node(5)

tree_Q = BinaryTree(q1)
tree_Q.add_left(q1, q2)
tree_Q.add_right(q1, q3)
tree_Q.add_left(q2, q4)
tree_Q.add_right(q2, q5)

assert isSameTree(tree_P.root, tree_Q.root) == True
```

<!-- grade_id: cell-e4a27208b29ac440, points: 0.5, grade: true, solution: false -->
```python
tree_A = BinaryTree(Node(1))
tree_A.add_left(tree_A.root, Node(2))
tree_A.add_right(tree_A.root, Node(3))

tree_B = BinaryTree(None)

assert isSameTree(tree_A.root, tree_B.root) == False
```

<!-- grade_id: cell-fdaae104c2710bed, points: 0.5, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL
```

<!-- grade_id: cell-92228ea749c04a75, points: 0.5, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL

```

# Problem 5: The Squirrel's Escape (2)
A little squirrel is stuck in a giant tree and needs to find the quickest way out! The tree is a binary tree, where each branch represents a path the squirrel can take.

To escape, the squirrel must start from the root and reach the nearest leaf node—a branch with no further splits. The goal is to find the shortest possible path (measured in nodes) from the starting position at the tree’s root to the closest leaf.

Example: This tree has a shortest path with length of 3
```
     1
   /   \
  2     3
 / \    /
4  5   6
  / \ /
  7 8 9

```
Shortest path: 1 --> 2 --> 4


Constraints:

- The number of nodes in the tree is in the range [0, 105].
- -1000 <= Node.val <= 1000


<!-- grade_id: cell-0d45cc585cf64233, points: 0, grade: false, solution: true -->
```python
def minDepth(root: BinaryTree):
    """
    :type root: Tree root
    :rtype: int
    """
    if not root:
        return 0

    if not root.left_child and not root.right_child:
        return 1
    

    if not root.left_child:
        return 1 + minDepth(root.right_child)
    elif not root.right_child: 
        return 1 + minDepth(root.left_child)
    else:
        return 1 + min(minDepth(root.left_child), minDepth(root.right_child))

```

<!-- grade_id: cell-3e902a3136ccd790, points: 0.5, grade: true, solution: false -->
```python
n1 = Node(1)
n2 = Node(2)
n3 = Node(3)
n4 = Node(4)
n5 = Node(5)
n6 = Node(6)
n7 = Node(7)
n8 = Node(8)
n9 = Node(9)

tree = BinaryTree(n1)

tree.add_left(n1, n2)
tree.add_right(n1, n3)
tree.add_left(n2, n4)
tree.add_right(n2, n5)
tree.add_left(n3, n6)

tree.add_left(n5, n7)
tree.add_right(n5, n8)
tree.add_left(n6, n9)

assert(minDepth(tree.root)) == 3
```

<!-- grade_id: cell-1d7a0b770001868c, points: 0.5, grade: true, solution: false -->
```python
n1 = Node(1)
n2 = Node(2)

tree = BinaryTree(n1)
tree.add_right(n1, n2)

assert(minDepth(tree.root)) == 2
```

<!-- grade_id: cell-a21c8e01870b8b21, points: 0.5, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL

```

<!-- grade_id: cell-10aef44684fd6538, points: 0.5, grade: true, solution: false -->
```python
# DO NOT DELETE THIS CELL

```